---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise. it's great

`@instructions`
1. first thing
2. Second thing
3) third thing

`@hint`
this is a cool hint bro

`@pre_exercise_code`
```{r}
a<-3
```

`@sample_code`
```{r}
# put 3 into a

	a<-3
```

`@solution`
```{r}
a<-3
```

`@sct`
```{r}
a<-3

success_msg("What a great job!")"
```
